## IPython Magic

IPython magic are special commands that can be used to interact with your System.

Enter the following into an IPython session:

```python
(an_integer, 
 a_list,
 a_dictionary,
 a_set,
 a_tuple) = 1, [1,2,3], {'k': 1}, {1,2,2,3}, (1,2)
```

### Writing a Function

Functions in python are written using the keyword `def`.

    def function_name(arg1, arg2):
        output_1 = do_something_with(arg1)
        output_2 = do_something_with(arg2)
        return 

```python
from sys import getsizeof

def sizeof_and_value(variable):
    return (getsizeof(variable), variable)
```

```python
sizeof_and_value(an_integer)
```

```python
sizeof_and_value(a_tuple)
```

### Write a Function

Write a function named `type_and_value` that returns the value and the type of a variable that is passed to it.

Use the block below to define the function. 

<!-- grade_id: cell-c48f82b1aae3c2f5, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

Make sure that your function can pass this test:

<!-- grade_id: cell-80a7fd4b5e0528b5, points: 2, grade: true, solution: false -->
```python
assert type_and_value(an_integer) == (int, 1)

```

## `%whos`

Prints a table with some basic details about each identifier you have defined interactively.


```python
%whos 
```

## Bash in IPython

Bash is a command line language used to interactive with an operating system.

Some simple Bash commands can be run in IPython/Jupyter, including

- `pwd`
- `cd`
- `ls`

### `pwd` - print working directory

```python
%pwd
```

### `cd` - change directory

```python
%cd src/
```

```python
%pwd
```

### `ls` - list files

```python
%ls
```

<!-- grade_id: cell-6a5baeb91a451736, points: 2, grade: true, solution: false -->
```python
# HIDDEN TEST

```

## IPython Magic commands

There are many IPython magic commands, but some of the more useful are

- `run`
- `matplotlib inline`
- `whos`


```python
%run a_simple_script.py
```

```python
%matplotlib inline
```

```python
%run a_simple_script.py
```

### Why didn't the image show up the first time?

<!-- grade_id: cell-e3c4593b0ce6604d, points: 3, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```