## The Python Numerical Stack

Consists of:

- numpy/scipy (vectors and computational mathematics)
- pandas (dataframes)
- matplotlib (plotting)
- seaborn (statistical plotting)
- scikit-learn (machine learning)

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import seaborn as sns

%matplotlib inline
```

**Note** We typically only import what we need from scikit-learn e.g.

```python
from sklearn.linear_model import LinearRegression
```

## The Iris Dataset

> This is perhaps the best known database to be found in the pattern recognition literature. Fisher's paper is a classic in the field and is referenced frequently to this day.
          — [UCI Machine Learning Repository](https://archive.ics.uci.edu/ml/datasets/iris)



```python
from sklearn.datasets import load_iris
```

> The Iris flower data set or Fisher's Iris data set is a multivariate data set introduced by the British statistician and biologist Ronald Fisher in his 1936 paper The use of multiple measurements in taxonomic problems as an example of linear discriminant analysis.
          — [Wikipedia](https://en.wikipedia.org/wiki/Iris_flower_data_set)



```python
IRIS = load_iris()
```

```python
type(IRIS.data)
```

```python
IRIS.data.shape
```

#### What does `.shape` do?

<!-- grade_id: cell-12f5b8407c965d99, points: 5, grade: true, solution: true -->
```python
# your code here
raise NotImplementedError
```

## Dataframes

We will load the Iris data into a dataframe for ease of manipulation.

```python
iris_df = pd.DataFrame(IRIS.data, columns=IRIS.feature_names)
```

```python
iris_df.head()
```


## Pair Plot

We will use Searborn to prepare a **Pair Plot** of the Iris dataset. A Pair Plot is an array of scatter plots, one for each pair of features in the data. Rather than plotting a feature against itself, the diagonal is rendered as a **probability distribution** of the given feature.


```python
sns.pairplot(iris_df)
```

## List Comprehension

We will use a **list comprehension** to remove the units and white space from the feature names to make them more "computer-friendly".


In general, list comprehensions have this form:

```python
lc = [do_something_to(var) for var in some_other_list]
```


```python
def square_number(x):
    return x**2
```

```python
[square_number(i) for i in (1,2,3,4,5)]
```

### Write your own list comprehension

Write a function that uses a list comprehension to change this list

    [1,2,3,4,5]
    
into this list

    [2,3,4,5,6]

<!-- grade_id: cell-2cdac04ef27891ff, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-efeded0deb1bb892, points: 3, grade: true, solution: false -->
```python
assert incr_list_by_1([1,2,3,4,5]) == [2,3,4,5,6]

```

### Remove Unit and White Space from Feature Name

Here we use a list comprehension to change the feature names:

```python
IRIS.feature_names
```

```python
iris_features_names = IRIS.feature_names
iris_features_names
```

```python
def remove_unit_and_white_space(feature_name):
    feature_name = feature_name.replace(' (cm)','')
    feature_name = feature_name.replace(' ', '_')
    return feature_name
```

```python
iris_features_names = [remove_unit_and_white_space(name) for name in iris_features_names]
```

```python
iris_features_names
```

```python
iris_df.columns = iris_features_names
iris_df.head()
```

## Export to CSV

Ultimately, we will export a CSV of the dataframe to disk. This will make it easy to access the same data from both Python and R.


```python
%ls
```

```python
%mkdir -p data
```

```python
%ls
```

```python
iris_df.to_csv('data/iris.csv')
```

```python

```